---
title: "DSS: CHAPTER 3"
tutorial:
  id: "03-population"
  version: 1.3
output:
  learnr::tutorial:
    progressive: true
    ace_theme: github
    theme: spacelab
    language:
      en:
        button:
          previoustopic: Previous Page
          nexttopic: Next Page
runtime: shiny_prerendered
description: "Review Exercises for Data Analysis for Social Science - CHAPTER 3."
---

```{r setup, include=FALSE}
library(gradethis)
library(learnr)
tutorial_options(exercise.checker = gradethis::grade_learnr, 
                 exercise.reveal_solution = TRUE,
                 exercise.error.check.code = "grade_code()")
knitr::opts_chunk$set(echo = FALSE)
chapter_id <- "Chapter 3"
num_tutorial <- 47
```

<style>

.tutorial-question .panel-body { padding: 6px 15px; }

.tutorial-question-container {
  margin-bottom: 10px;
}

.borderexample {
 border-width:1px;
 border-style:solid;
 border-color:dodgerblue;
 padding: 10px;
}
</style>

<style>
table, td, th {
  border: 1px solid lightgray;
}

table {
  border-collapse: collapse;
  width: 100%;
  margin-bottom: 20px;
  margin-top: 15px;
}

td { text-align: center;}
th { text-align: center;}
th, td { padding: 15px; }
</style>


## <span style=color:white> OVERVIEW

<div style="overflow: auto;">
<img src="images/dss.png" 
     width="200" 
     height="250" 
     style="float: left; margin-right: 20px; margin-bottom: 10px;" />
     
***What are these tutorials and how to use them?*** These tutorials are meant for students using DSS: <a href="https://ellaudet.github.io/dss_book/" style="color: inherit;">Llaudet and Imai's <i>Data Analysis for Social Science: A Friendly and Practical Introduction</i></a>.

They are designed to be completed *after* reading and performing the analyses explained in the book — they are meant to supplement, not replace, the book.
Each includes multiple-choice questions and hands-on coding exercises, with instant feedback so students can check their understanding on their own. At the end, students can generate a progress report of their work. The questions are organized into topics, 
each covering a handful of pages from the book that can be read in one sitting. Here is the recommended sequence:

<table>
  <tr>
    <th><span style=color:lightgray>$\blacktriangleright$ </span><span style=color:gray>  STEP 1  <span style=color:lightgray> $\blacktriangleright$ </span></th>
    <th><span style=color:lightgray> $\blacktriangleright$  </span><span style=color:gray>  STEP 2  <span style=color:lightgray> $\blacktriangleright$ </span></th>
    <th> <span style=color:lightgray> $\blacktriangleright$  </span><span style=color:gray>  STEP 3 <span style=color:lightgray> $\blacktriangleright$</span></th>
    <th> <span style=color:lightgray> $\blacktriangleright$  </span><span style=color:gray>  STEP 4 <span style=color:lightgray> $\blacktriangleright$</span></th>
  </tr>
  <tr>
<td> read the section(s) for a topic </td>
<td> perform the analyses on your computer </td>
<td> complete the review exercises </td>
<td> move on to the next topic </td>
  </tr>
</table>
</div>

***What does this particular tutorial cover?*** This is the tutorial for chapter 3, where we learn how to infer population characteristics from a representative sample, how to handle missing data, and how to visualize and summarize both the distribution of a single variable and the relationship between two variables. The review exercises for this chapter are organized into two topics:

<table>
  <tr>
    <th><span style=color:gray> Topic </span></th>
    <th><span style=color:gray> Readings </span></th>
    <th><span style=color:gray> Key Concepts and Notation </span></th>
    <th><span style=color:gray> R Functions and Operators </span></th>
  </tr>
  <tr>
    <td> Survey Research and Data Exploration  </td>
    <td> 3-3.4 </td>
    <td> sample, representative sample, random sampling, population, sample size (n), N, sampling frame, unit nonresponse, item nonresponse, misreporting, table of frequencies, table of proportions, missing values (NA), listwise deletion, two-way table of frequencies, two-way table of proportions, histogram, density histogram, descriptive statistics, mean, median, standard deviation, variance </td>
    <td> table(), prop.table(), na.omit(), hist(), mean(), median(), sd(), var() </td>
  </tr>
  <tr>
    <td> Exploring the Relationship Between Two Variables </td>
    <td> 3.5-3.6 </td>
    <td> scatter plot, correlation coefficient, direction and strength of linear association </td>
    <td> plot(), abline(), cor() </td>
  </tr>
</table>

Let's start! <br>

## <span style=color:dodgerblue> 1. Survey Research and Data Exploration
The review exercises in this topic cover:

<table>
  <tr>
    <th> Readings </th>
    <th> Key Concepts and Notation </th>
    <th> R Functions and Operators </th>
  </tr>
  <tr>
    <td> 3-3.4 </td>
    <td> sample, representative sample, random sampling, population, sample size (n), N, sampling frame, unit nonresponse, item nonresponse, misreporting, table of frequencies, table of proportions, missing values (NA), listwise deletion, two-way table of frequencies, two-way table of proportions, histogram, density histogram, descriptive statistics, mean, median, standard deviation, variance </td>
    <td> table(), prop.table(), na.omit(), hist(), mean(), median(), sd(), var() </td>
  </tr>
</table>

## <font size="+2"> <span style=color:dodgerblue> $\textrm{             }$$\textrm{             }$$\textrm{             }$ Multiple-Choice Questions </span></font>

<div class="borderexample">
Instructions:
<ul style="padding-left: 0.3em; margin-left: 0; list-style-position: inside;">
  <li>Select correct answer(s) and click <i>Submit Answer</i></li>
  <li>If incorrect, click <i>Try Again</i> until correct</li>
  <li>Progress is saved automatically; click <i>Start Over</i> to delete all your answers</li>
</ul>
</div>

```{r ch3_topic1_mc}
qnum <- 0
quiz(caption=NULL,
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Which of the three goals of quantitative social science research is chapter 3 about?"),
    answer("measuring a quantity of interest", correct = TRUE),
    answer("explaining a quantity of interest", message = "Chapters 2 and 5 are about explaining."),
    answer("predicting a quantity of interest", message = "Chapter 4 is about predicting."),
    post_message = "Chapter 3 is about measuring a quantity of interest; specifically, about inferring the characteristics of a population from a sample.<br><br> See DSS §1.1 and the introduction to chapter 3.",
    random_answer_order = TRUE,
    allow_retry = TRUE
),
     
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> In the British Election Study (BES), researchers surveyed just under 31,000 people to infer the attitudes of more than 46 million eligible UK voters. In that study..."),
    answer("*n* = 31,000 and *N* = 46 million", correct = TRUE),
    answer("*n* = 46 million and *N* = 31,000"),
    answer("*n* = 31,000 and *N* = 31,000"),
    answer("*n* = 46 million and *N* = 46 million"),
    post_message = "*n* is the sample size and *N* is the population size. Here, the 31,000 people surveyed are the sample (*n*); the 46 million eligible voters they are meant to represent are the population (*N*).<br><br> See DSS §3.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),
  
    question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> When using a sample to infer population characteristics, it is important that:"),
    answer("The sample be representative of the population of interest", correct = TRUE),
    answer("The treatment and control groups be comparable"),
    answer("The dataset contains both the treatment and the outcome variables"),
    answer("The sample includes every individual from the population of interest"),
    post_message = "Only if the sample is representative of the population of interest can we use it to infer that population's characteristics.<br><br> See DSS §3.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),
  
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The best way to draw a representative sample from a population is:"),
    answer("To randomly select individuals from the population of interest", correct = TRUE),
    answer("To assign treatment at random"),
    answer("To select individuals from the population of interest who are easily reachable"),
    answer("To allow individuals from the population of interest to choose whether they want to be part of the sample"),
    post_message = "Random sampling gives every individual in the population the same chance of being selected, which makes the sample representative of the population, on average.<br><br> See DSS §3.2.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
),
  
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, 
  " - SELECT ALL THAT APPLY:</span> Random sampling..."),
  answer("determines who is selected for the sample", correct = TRUE),
  answer("determines who receives the treatment", message="You are thinking about random treatment assignment."),
  answer("makes the treatment and control groups comparable, on average", message="You are thinking about random treatment assignment."),
  answer("makes the sample representative of the population, on average", correct = TRUE),
  answer("enables valid estimates of average causal effects", message="You are thinking about random treatment assignment."),
  answer("enables valid inferences about population characteristics", correct = TRUE),
  post_message = "See DSS §3.2.1 and §2.4.1.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, 
  ":</span> Are students at a university in Massachusetts (USA) representative of all residents in that state?"),
  answer("No", correct = TRUE),
  answer("Yes"),
  post_message = "The sample would not represent children, seniors, or residents without a college education, among many other groups.<br><br> See DSS §3.2.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Suppose we want to know what proportion of students at your university intend to vote in the next election, so we survey the 200 students in one political science class. What is the main problem with this approach?"), 
  answer("The sample is unlikely to be representative of the target population", correct = TRUE),
  answer("The sample size is likely to be too small"),
  answer("We did not use random treatment assignment", message="Random treatment assignment has to do with how we assign treatment, not with how we select the sample; and there is no treatment here."),
  answer("There is no problem with this approach"),
  post_message = "The students in one political science class are unlikely to resemble the university's student body as a whole (among other things, they are likely to be more interested in politics and thus more likely to vote), so the sample is unlikely to be representative of the target population.<br><br> See DSS §3.2.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Suppose instead that we email *all the students at the university* and collect data from those who choose to respond. What is the main problem with this approach?"),
    answer("The sample is unlikely to be representative of the target population", correct = TRUE),
    answer("The sample size is likely to be too small"),
    answer("Emails cannot be used to conduct surveys"),
    answer("There is no problem with this approach because all students were contacted"),
    post_message = "The students who choose to respond may differ systematically from those who choose not to respond, making the resulting sample unrepresentative.<br><br> See DSS §3.2.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> If X = {0, 1, 0, 0, 1}, the table of frequencies of X is..."),
  answer("0&nbsp;&nbsp;1<br>3&nbsp;&nbsp;2", correct = TRUE),
  answer("0&nbsp;&nbsp;1<br>2&nbsp;&nbsp;3"),
  answer("0&nbsp;&nbsp;1<br>5&nbsp;&nbsp;5"),
  answer("0&nbsp;&nbsp;1<br>1&nbsp;&nbsp;1"),
  post_message = "A table of frequencies shows the values the variable takes along with the number of observations that take each value. Here, X contains three 0s and two 1s, so the frequencies are 3 and 2.<br><br> See DSS §3.3.2.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> And, the table of proportions of X is..."),
  answer("0&nbsp;&nbsp;&nbsp;1<br>0.6&nbsp;0.4", correct = TRUE),
  answer("0&nbsp;&nbsp;&nbsp;1<br>3&nbsp;&nbsp;&nbsp;2"),
  answer("0&nbsp;&nbsp;&nbsp;1<br>0.4&nbsp;0.6"),
  answer("0&nbsp;&nbsp;&nbsp;1<br>0.5&nbsp;0.5"),
  post_message = "X contains three 0s and two 1s out of five observations, so the proportions are 3/5 = 0.6 and 2/5 = 0.4.<br><br> See DSS §3.3.3.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> To explore the contents of a single variable, we can..."),
    answer("Compute its table of frequencies or of proportions", correct = TRUE),
    answer("Compute its two-way table of frequencies or of proportions", 
           message = "Two-way tables require two variables."),
    answer("Create a scatter plot to visualize its distribution", 
           message = "Scatter plots visualize the relationship between two variables, not the distribution of a single variable."),
    answer("Create a histogram to visualize its distribution", correct = TRUE),
    answer("Compute its mean or median to summarize the center of the distribution", correct = TRUE),
    answer("Compute the correlation coefficient to summarize the center of the distribution"),
    answer("Compute its standard deviation or variance to summarize the spread of the distribution", correct = TRUE),
    post_message = "See DSS §3.6.",
    random_answer_order = FALSE,
    allow_retry = TRUE
)
)
```

```{r echo=FALSE}
set.seed(123)
x <- rnorm(300, mean = 50, sd = 12)
hist(x,
     main = "",
     xlab = "",
     ylab = "Frequency",
     col = "steelblue",
     border = "white")
mtext("X", side = 1, line = 3, adj = 1)
```

```{r ch3_topic1_mc_graphs}
quiz(caption=NULL,

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The type of graph displayed above is called a..."),
    answer("Histogram", correct = TRUE),
    answer("Scatter plot"),
    answer("Table of frequencies"),
    answer("Density plot"),
    post_message = "See DSS §3.4.4.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),
  
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Looking at a histogram, we can get a sense of:"),
  answer("The distribution of the variable", correct = TRUE),
  answer("The center of the distribution", correct = TRUE),
  answer("The spread of the distribution", correct = TRUE),
  answer("The range of values the variable takes", correct = TRUE),
  answer("The relationship between two variables", message="You might be thinking about a scatter plot."),
  answer("The direction of the linear association between two variables", message="You might be thinking about a scatter plot."),
  answer("The strength of the linear association between two variables", message="You might be thinking about a scatter plot."),
  post_message = "See DSS §3.4.4.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Which are true?"),
  answer("In a histogram, the y-axis displays frequencies (counts), while in a density histogram the y-axis displays densities", correct = TRUE),
  answer("Density histograms are more useful than regular histograms when comparing distributions with substantially different numbers of observations", correct = TRUE),
  answer("When the width of the bins is constant, the relative heights of the bins in a density histogram reflect the relative proportions of observations in each bin", correct = TRUE),
  post_message = "See DSS §3.4.5.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> If X = {10, 5, 8, 12, 15}, what are the mean and the median of X?"),
  answer("Mean = 10 and Median = 10", correct = TRUE),
  answer("Mean = 8 and Median = 10"),
  answer("Mean = 10 and Median = 8"),
  answer("Mean = 12 and Median = 10"),
  post_message = "The mean is (10+5+8+12+15)/5 = 50/5 = 10. Sorted, X = {5, 8, 10, 12, 15}, so the middle value — the median — is also 10.<br><br> See DSS §1.8.2 and §3.4.6.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> If X = {1, 0, 1, 0, 0}, what are the mean and the median of X?"),
  answer("Mean = 0.4 and Median = 0", correct = TRUE),
  answer("Mean = 0 and Median = 0.4"),
  answer("Mean = 0.5 and Median = 0"),
  answer("Mean = 0.4 and Median = 1"),
  post_message = "The mean is (1+0+1+0+0)/5 = 2/5 = 0.4. Sorted, X = {0, 0, 0, 1, 1}, so the middle value — the median — is 0.<br><br> See DSS §1.8.2 and §3.4.6.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> If X is a binary variable that identifies students who voted in the last election and the mean of X is 0.72, it means that..."),
  answer("72% of students voted in the last election", correct = TRUE),
  answer("On average, 72% of students voted in the last election", message="The mean of a binary variable represents a proportion, not an average, so the word \"average\" should not appear in the interpretation."),
  answer("72 students voted in the last election", message="The mean of a binary variable represents a proportion of the observations, not a count of them."),
  answer("The average student voted 0.72 times in the last election", message="X identifies whether a student voted, not how many times."),
  post_message = "The mean of a binary variable represents the proportion of the observations that have the characteristic identified by the variable, which we express as a percentage after multiplying by 100 (0.72×100 = 72%). The word \"average\" should not appear in the interpretation.<br><br> See DSS §1.8.2.",
  random_answer_order = TRUE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The standard deviation of a variable roughly measures..."),
    answer("the average distance between the observations and the mean", correct = TRUE),
    answer("the distance between the maximum and minimum values of the variable"),
    answer("the proportion of observations that fall within one unit of the mean"),
    post_message = "The standard deviation summarizes the spread of a distribution: roughly, how far the observations lie from the mean, on average.<br><br> See DSS §3.4.6.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  )
)
```

```{r echo=FALSE, fig.height=2.5}
par(mar = c(1, 0.5, 1, 1))
x <- seq(0, 100, length.out = 500)
plot(x, dnorm(x, mean = 50, sd = 8),
     type = "l", lty = 1, lwd = 2, col = "steelblue",
     axes = FALSE, ann = FALSE)
lines(x, dnorm(x, mean = 50, sd = 20),
      lty = 2, lwd = 2, col = "steelblue")
```

```{r ch3_topic1_mc_graphs2}
quiz(caption=NULL,
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Looking at the two distributions above, we can see that..."),
    answer("the solid distribution has the larger standard deviation", message="A larger standard deviation means the observations are more spread out, which makes the distribution wider and flatter — not taller and narrower."),
    answer("the dashed distribution has the larger standard deviation", correct = TRUE),
    answer("in the dashed distribution, the observations are farther from the mean, on average", correct = TRUE),
    answer("in the solid distribution, the observations are farther from the mean, on average", message="In the solid distribution the observations are tightly concentrated around the mean, so their average distance from it is smaller."),
    answer("the solid distribution is flatter and more spread out", message="The solid distribution is the taller, narrower one."),
    answer("the dashed distribution is flatter and more spread out", correct = TRUE),
    post_message = "See DSS §3.4.6.",
    random_answer_order = FALSE,
    allow_retry = TRUE
),
     
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> If the standard deviation of X is 25 miles, then the variance of X is..."),
    answer("625 miles²", correct = TRUE),
    answer("5 miles"),
    answer("625 miles"),
    answer("5 miles²"),
    answer("impossible to know with the information given"),
    post_message = "The variance is the square of the standard deviation: 25² = 625, and miles × miles = miles².<br><br> See DSS §3.4.6.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),
  
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> R represents a missing value as ..."),
    answer("`NA`", correct = TRUE),
    answer("`NULL`"),
    answer("`0`"),
    answer("`missing`"),
    post_message = "See DSS §3.4.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> In R, what does `na.omit()` do?"),
    answer("It removes from the dataframe all observations that contain at least one missing value", correct = TRUE),
    answer("It replaces all missing values in the dataframe with zeros"),
    answer("It removes from the dataframe all variables that contain at least one missing value"),
    answer("It counts the number of missing values in the dataframe"),
    post_message = "`na.omit()` performs listwise deletion: it removes from a dataframe every observation containing at least one NA.<br><br> See DSS §3.4.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> To create the table of proportions of a variable, we run..."),
    answer("`prop.table(table(data$variable))`", correct = TRUE),
    answer("`table(data$variable)`", message="This is the code to create a table of frequencies."),
    answer("`prop.table(data$variable)`", message="Hint: `prop.table()` converts a table of frequencies into a table of proportions."),
    post_message = "`prop.table()` converts a table of frequencies into a table of proportions, so we specify as its main argument the code producing the table of frequencies using `table()`.<br><br> See DSS §3.3.3.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> To compute the mean of *age* stored in `bes1`, for ONLY the observations where *leave* equals 0, we run..."),
    answer("`mean(bes1$age[bes1$leave==0])`", correct = TRUE),
    answer("`mean(bes1$age[bes1$leave=0])`", message = "The operator we use to set the logical test is `==`, not `=`."),
    answer("`mean(age[leave==0])`", message = "This code does not identify the object where the variables are stored."),
    answer("`mean(bes1$leave[bes1$age==0])`", message = "This code would compute the mean of *leave* for observations where *age* equals 0, not the other way around."),
    post_message = "To the left of `[]` we specify the variable we want to subset: `bes1$age`; inside the brackets, the logical test that selects the observations: `bes1$leave==0`. <br><br> See DSS §2.5.3 and §3.4.6.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  )
)
```


## <font size="+2"> <span style=color:dodgerblue> $\textrm{             }$$\textrm{             }$$\textrm{             }$ Coding Exercises </span></font>

<div class="borderexample">
Instructions:
<ul style="padding-left: 0.3em; margin-left: 0; list-style-position: inside;">
<li>Type code in the <i>R code</i> window</li>
  <li>Click <i>Run Code</i> to see the output</li>
  <li>Click <i>Submit Answer</i> when satisfied</li>
  <li>If incorrect, modify your code and try again until correct</li>
  <li>Click <i>Hint</i> as needed</li>
  <li>Progress is saved automatically; click <i>Start Over</i> to delete all your answers</li>
</ul>
</div>

<br><b>For the next few questions, we are going to work with the BES dataset, where the unit of observation is survey respondents and the variables are as described below:</b>

<table>
  <tr> 
    <th> variable </th>
    <th> description </th>
  </tr>
  <tr>
    <td> *vote* </td>
    <td> respondent's vote intention in the EU referendum: "leave", "stay", "don't know", or "won't vote" </td>
  </tr>
  <tr>  
    <td>*leave* </td>
    <td> identifies leave voters: 1 = intends to vote "leave" or 0 = intends to vote "stay"; <br>NA = either "don't know" or "won't vote" </td>
  </tr>
  <tr>  
    <td> *education* </td>
    <td> respondent's highest educational qualification: 1 = no qualifications, 2 = GCSE, 3 = GCE A level, <br>4 = undergraduate degree, 5 = postgraduate degree; NA = no answer </td>
  </tr>
  <tr>
    <td>*age* </td>
    <td> respondent's age in years </td>
  </tr>
</table>

<br><b>In the following exercises, we analyze the BES survey data to answer: *Who supported Brexit?*</b> <span style=color:dodgerblue> **SETUP:**</span> **Suppose we have already set the working directory and run the code below to load the data, remove observations with missing values, store the new dataframe as `bes1`, and ask R to show us the first six observations.**

```{r bes, echo=FALSE}
bes <- read.csv("data/BES.csv")
bes1 <- na.omit(bes)
```

```{r bes_FALSE2, eval=FALSE, echo=TRUE}
bes <- read.csv("BES.csv") # reads and stores data
bes1 <- na.omit(bes) # removes observations with NAs and stores new dataframe
```

```{r bes_render, include=FALSE}
bes <- read.csv("data/BES.csv")
bes1 <- na.omit(bes)
```

```{r bes_head, echo=TRUE}
head(bes1) # shows the first six observations
```

<br> **Use the dataframe with no missing data, `bes1`, for all the exercises.**

<span style=color:dodgerblue> **Q`r qnum <- qnum+1; qnum`:** </span> **Create the histogram of *age* for Brexit non-supporters (leave=0).**

```{r ch3_topic2_r8, exercise=TRUE, exercise.setup = "bes"}

```

<div id="ch3_topic2_r8-hint">
Hint: Use `hist()`. Subset *age* to include only observations where *leave* equals 0 using `bes1$age[bes1$leave==0]`. See DSS §3.4.4.
</div>

```{r ch3_topic2_r8-solution}
hist(bes1$age[bes1$leave==0])
```

```{r ch3_topic2_r8-check}
grade_code("<br><br> See DSS §3.4.4.")
```

<br> <span style=color:dodgerblue> **Q`r qnum <- qnum+1; qnum`:** </span> **Create the histogram of *age* for Brexit supporters (leave=1).**

```{r ch3_topic2_r9, exercise=TRUE, exercise.setup = "bes"}

```

<div id="ch3_topic2_r9-hint">
Hint: Use `hist()`. Subset *age* to include only observations where *leave* equals 1 using `bes1$age[bes1$leave==1]`. See DSS §3.4.4.
</div>

```{r ch3_topic2_r9-solution}
hist(bes1$age[bes1$leave==1])
```

```{r ch3_topic2_r9-check}
grade_code("<br><br> See DSS §3.4.4.")
```

```{r ch3_topic2_r8_inter}
quiz(caption=NULL,
  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Comparing the two histograms above, we see that in the BES sample:"),
  answer("Brexit supporters tended to be older than non-supporters", correct = TRUE),
  answer("Brexit supporters tended to be younger than non-supporters"),
  answer("The age distributions of supporters and non-supporters were very similar"),
  post_message = "See DSS §3.4.4.",
  random_answer_order = TRUE,
  allow_retry = TRUE
  )
)
```

<br> <span style=color:dodgerblue> **Q`r qnum <- qnum+1; qnum`:** </span> **Compute the mean, median, and standard deviation (in this order) of *age* for Brexit non-supporters (leave=0) and then again for Brexit supporters (leave=1).**

```{r ch3_topic2_r7, exercise=TRUE, exercise.setup = "bes"}

```

<div id="ch3_topic2_r7-hint">
Hint: For the first three, subset *age* to include only the observations for which *leave* equals 0, using `bes1$age[bes1$leave==0]` as the argument for `mean()`, `median()`, and `sd()`. Then repeat with `bes1$leave==1`. See DSS §3.4.6.
</div>

```{r ch3_topic2_r7-solution}
mean(bes1$age[bes1$leave==0])
median(bes1$age[bes1$leave==0])
sd(bes1$age[bes1$leave==0])
mean(bes1$age[bes1$leave==1])
median(bes1$age[bes1$leave==1])
sd(bes1$age[bes1$leave==1])
```

```{r ch3_topic2_r7-check}
grade_code("<br><br> See DSS §3.4.6.")
```

```{r ch3_topic2_r7_inter}
quiz(caption=NULL,
  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> This means that, in the BES sample: 
<br><br> Brexit non-supporters were about... years old on average, half of them were ... years old or younger, and their ages were about ... years from the mean, on average
<br><br> Brexit supporters were about... years old on average, half of them were ... years old or younger, and their ages were about ... years from the mean, on average"),
    answer("47, 48, and 17; and 55, 58, and 15", correct = TRUE),
    answer("48, 47, and 17; and 58, 55, and 15", message="Check the order."),
    answer("17, 47, and 48; and 15, 55, and 58", message="Check the order."),
    answer("55, 58, and 15; and 47, 48, and 17", message="These are the values for supporters first and non-supporters second, but the question asks about non-supporters first."),
    post_message = "Among non-supporters: the mean age is about 47 years, the median is 48 years (so half were 48 or younger), and the standard deviation is about 17 years (so individual ages were about 17 years from the mean, on average). Among supporters, the same three statistics are about 55, 58, and 15 years. All three are measured in years, the unit of measurement of *age*.<br><br> See DSS §3.4.6.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),
  
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> From this analysis we conclude that: Assuming the BES sample was representative of all eligible UK voters at the time, we infer that Brexit supporters tended to be..."),
    answer("older than non-supporters", correct = TRUE),
    answer("younger than non-supporters", message="Check the direction: the mean and median ages are higher among supporters."),
    post_message = "The mean age is higher among supporters than among non-supporters. We start the conclusion statement with the assumption that the sample is representative because the validity of our inference depends on it.<br><br> See DSS §3.4.6.",
    random_answer_order = TRUE,
    allow_retry = TRUE
)
  )
```


## <span style=color:dodgerblue> 2. Exploring the Relationship Between Two Variables
The review exercises in this topic cover:

<table>
  <tr>
    <th> Readings </th>
    <th> Key Concepts and Notation </th>
    <th> R Functions and Operators </th>
  </tr>
  <tr>
    <td> 3.5-3.6 </td>
    <td> scatter plot, correlation coefficient, direction and strength of linear association </td>
    <td> plot(), abline(), cor() </td>
  </tr>
</table>

## <font size="+2"> <span style=color:dodgerblue> $\textrm{             }$$\textrm{             }$$\textrm{             }$ Multiple-Choice Questions </span></font>

<div class="borderexample">
Instructions:
<ul style="padding-left: 0.3em; margin-left: 0; list-style-position: inside;">
  <li>Select correct answer(s) and click <i>Submit Answer</i></li>
  <li>If incorrect, click <i>Try Again</i> until correct</li>
  <li>Progress is saved automatically; click <i>Start Over</i> to delete all your answers</li>
</ul>
</div>

```{r ch3_topic_pre_mc}
quiz(caption=NULL,
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> To visualize..."),
    answer("the distribution of one variable, we can create a histogram", correct=TRUE),
    answer("the relationship between two variables, we can create a scatter plot", correct = TRUE),
    answer("the distribution of one variable, we can create a scatter plot"),
    answer("the relationship between two variables, we can create a histogram"),
    post_message = "See DSS §3.4.4 and §3.5.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  )
)
```

```{r echo=FALSE}
set.seed(123)
x <- rnorm(300, mean = 50, sd = 12)
y <- 2 + 3*x + rnorm(300, mean=0, sd=10)
plot(x, y,
     main = "",
     xlab = "",
     ylab = "",
     pch = 19,
     col = "steelblue",
     cex = 0.7)
mtext("X", side = 1, line = 3, adj = 1)
mtext("Y", at = 260, side = 2, line = 3, adj = 1, las=1)
```

```{r ch3_topic3_mc}
quiz(caption=NULL,
     
question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The type of graph displayed above is called a..."),
    answer("Histogram"),
    answer("Scatter plot", correct = TRUE),
    answer("Table of frequencies"),
    answer("Density plot"),
    post_message = "A scatter plot displays the relationship between two variables: X along the x-axis, Y along the y-axis, and one dot per observation positioned according to its values of X and Y.<br><br> See DSS §3.5.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
       ),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> In the scatter plot above:"),
  answer("The direction of the linear association between the two variables is positive", correct = TRUE),
  answer("The direction of the linear association between the two variables is negative"),
  answer("The linear association between the two variables is relatively weak"),
  answer("The linear association between the two variables is relatively strong", correct = TRUE),
  answer("The relationship between X and Y is causal", message="Sadly, this is not something we can see in a scatter plot."),
  answer("The relationship between X and Y is NOT causal", message="Sadly, this is not something we can see in a scatter plot."),
  post_message = "The slope of the (currently imaginary) line of best fit is positive (it tilts upward from left to right), so the linear association is positive. The dots cluster tightly around the line of best fit, so the linear association is relatively strong. Whether the relationship between X and Y is causal or not is not something that we can observe in a scatter plot.<br><br> See DSS §3.5.1.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Does it matter which variable is plotted on the x-axis and which on the y-axis?"),
  answer("Yes, if our purpose is simply to observe the direction and strength of the linear association between X and Y", message="The direction and strength of a linear association are the same whichever variable goes on which axis: swapping them flips the plot without changing what it tells us."),
  answer("No, if our purpose is simply to observe the direction and strength of the linear association between X and Y", correct = TRUE),
  answer("Yes, if our purpose is to observe the relationship between two variables for making predictions or estimating causal effects", correct = TRUE),
  answer("No, if our purpose is to observe the relationship between two variables for making predictions or estimating causal effects", message="As we will see in chapter 4, when making predictions, the predictor goes on the x-axis and the outcome variable on the y-axis. And as we will see in chapter 5, when estimating causal effects, the treatment variable goes on the x-axis and the outcome variable of interest on the y-axis."),
  post_message = "In chapter 3, our only purpose is to observe the direction and strength of the relationship between two variables, so it does not matter which one is plotted on which axis. However, as we will see in chapter 4, when making predictions, the predictor goes on the x-axis and the outcome variable on the y-axis; and as we will see in chapter 5, when estimating causal effects, the treatment variable goes on the x-axis and the outcome variable of interest on the y-axis.<br><br> See DSS §3.5.1.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The correlation between X and Y can range from:"),
    answer("–1 to 1", correct = TRUE),
    answer("0 to 1"),
    answer("–100 to 100"),
    answer("It depends on the units of measurement of X and Y"),
    post_message = "The correlation coefficient always falls between –1 and 1, whatever the units of X and Y.<br><br> See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> From the scatter plot above, we expect..."),
  answer("cor(X,Y) > 0 and |cor(X,Y)| to be closer to 1 than to 0", correct = TRUE),
  answer("cor(X,Y) < 0 and |cor(X,Y)| to be closer to 1 than to 0", message="The slope of the line of best fit is positive so the correlation is positive, not negative."),
  answer("cor(X,Y) > 0 and |cor(X,Y)| to be closer to 0 than to 1", message="The dots are closely clustered around the line of best fit, so the linear association between two variables is relatively strong, not weak"),
  post_message = "Since the linear association is positive, the correlation is positive. Since the linear association is strong, the absolute value of the correlation is closer to 1 than to 0.<br><br> See DSS §3.5.2.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),

question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> From the scatter plot above, we infer that:"),
  answer("Higher values of X are typically associated with higher values of Y, on average", correct=TRUE),
  answer("Higher values of X are typically associated with lower values of Y, on average", message="The ilnear association is positive, not negative."),
  answer("On average, either increases in X cause increases in Y or increases in Y cause increases in X", message="Correlation does not necessarily imply causation"),
    answer("On average, either increases in X cause decreases in Y or increases in Y cause decreases in X", message="Correlation does not necessarily imply causation"),
  post_message = "The association is positive, so higher values of X are typically associated with higher values of Y. We cannot infer anything about the causal relationship between the variables.<br><br> See DSS §3.5.1.",
  random_answer_order = FALSE,
  allow_retry = TRUE
),


question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> The correlation between outside temperature and heating bills is..."),
    answer("Negative, because higher outside temperatures are typically associated with lower heating bills", correct = TRUE),
    answer("Positive, because higher outside temperatures are typically associated with higher heating bills"),
    answer("Zero, because temperature and heating bills are not related in any way"),
    answer("It depends on the season"),
    post_message = "The two variables typically move in opposite directions and thus the correlation is negative.<br><br> See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Which correlation indicates a stronger linear association: –0.65 or 0.3?"),
    answer("–0.65, because the absolute value of the correlation is larger", correct = TRUE),
    answer("0.3, because the correlation is positive"),
    answer("They indicate an equally strong linear association because both correlations are between –1 and 1"),
    answer("We cannot compare the strength of the linear associations without knowing the units of measurement of the variables"),
    post_message = "The sign gives the direction of the association; the absolute value gives its strength. Since |–0.65| = 0.65 > |0.3| = 0.3, the –0.65 association is the stronger one.<br><br> See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Suppose the correlation between shoe size and reading ability among elementary school children is 0.8. Could we infer that having bigger feet causes better reading skills, on average?"),
    answer("No, correlation does not necessarily imply causation", correct = TRUE),
    answer("Yes, a correlation of 0.8 is positive and strong enough to conclude that bigger feet cause better reading"),
    answer("Yes, but only if we have a large enough sample size"),
    answer("No, a correlation of 0.8 is too weak to draw any conclusions"),
    post_message = "Correlation does not necessarily imply causation. A third variable, such as age, could be driving both: older children tend to have both larger feet and stronger reading skills.<br><br> See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> Which is true?"),
    answer("cor(Y,X) = cor(X,Y)", correct = TRUE),
    answer("A steeper line of best fit always means a higher correlation in absolute terms", message = "What determines the strength of the linear association (the absolute value of the correlation coefficient) is how tightly the dots cluster around the line, not the slope of the line."),
    answer("If cor(Y,X) = 0, X and Y have absolutely no relationship with each other", message = "A correlation of zero means there is no *linear* relationship between the variables, but a nonlinear relationship (such as a parabolic one) could still exist."),
    post_message = "See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  ),

  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> To create the scatter plot of X and Y (with X on the x-axis and Y on the y-axis), we could run..."),
    answer("`plot(X,Y)`", correct = TRUE),
    answer("`plot(y=Y, x=X)`", correct = TRUE),
    answer("`plot(Y,X)`"),
    answer("`plot(x=X, y=Y)`", correct = TRUE),
    answer("`plot(x=Y, y=X)`"),
    post_message = "If we do not use the names of the arguments, `plot()` plots the first variable on the x-axis and the second on the y-axis. If we use the names of the arguments `x=` and `y=`, the order no longer matters.<br><br> See DSS §3.5.1.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  )
)
```

## <font size="+2"> <span style=color:dodgerblue> $\textrm{             }$$\textrm{             }$$\textrm{             }$ Coding Exercises </span></font>

<div class="borderexample">
Instructions:
<ul style="padding-left: 0.3em; margin-left: 0; list-style-position: inside;">
<li>Type code in the <i>R code</i> window</li>
  <li>Click <i>Run Code</i> to see the output</li>
  <li>Click <i>Submit Answer</i> when satisfied</li>
  <li>If incorrect, modify your code and try again until correct</li>
  <li>Click <i>Hint</i> as needed</li>
  <li>Progress is saved automatically; click <i>Start Over</i> to delete all your answers</li>
</ul>
</div>

<br><b>For the next few questions, we work with the UK district-level data, where the unit of observation is UK districts and the variables are as described below:</b>

<table>
  <tr> 
    <th> variable </th>
    <th> description </th>
  </tr>
  <tr>
    <td> *name* </td>
    <td> name of the district </td>
  </tr>
  <tr>  
    <td>*leave* </td>
    <td> vote share received by the leave camp in the district (in percentages) </td>
  </tr>
  <tr>  
    <td> *high_education* </td>
    <td> proportion of district's residents with an undergraduate degree, professional qualification, or equivalent (in percentages) </td>
  </tr>
</table>

```{r dis, echo=FALSE}
dis <- read.csv("data/UK_districts.csv")
dis1 <- na.omit(dis)
```

<br> <span style=color:dodgerblue> **SETUP: **</span> **Suppose we have already set the working directory and run the code below to load the data, remove observations with missing values, store the new dataframe as `dis1`, and ask R to show us the first six observations.**

```{r dis_FALSE2, eval=FALSE, echo=TRUE}
dis <- read.csv("UK_districts.csv") # reads and stores data
dis1 <- na.omit(dis) # removes observations with NAs and stores new dataframe
```

```{r dis_render, include=FALSE}
dis <- read.csv("data/UK_districts.csv")
dis1 <- na.omit(dis)
```

```{r dis_head, echo=TRUE}
head(dis1) # shows the first six observations
```

<br> **Use the dataframe with no missing data, `dis1`, for all the exercises.**

<span style=color:dodgerblue> **Q`r qnum <- qnum+1; qnum`:** </span> **Use `plot()` to create a scatter plot of *high_education* (on the x-axis) and *leave* (on the y-axis).**

```{r ch3_topic3_r1, exercise=TRUE, exercise.setup = "dis"}

```

<div id="ch3_topic3_r1-hint">
Hint: `plot()` requires two arguments: the code identifying the X variable and the code identifying the Y variable (in that order), or use the named arguments `x=` and `y=`. See DSS §3.5.1.
</div>

```{r ch3_topic3_r1-solution}
plot(x=dis1$high_education, y=dis1$leave)
```

```{r ch3_topic3_r1-check}
grade_code("<br><br> See DSS §3.5.1.")
```

```{r ch3_topic3_r1_inter}
quiz(caption=NULL,
  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Based on the scatter plot of *high_education* and *leave*, their linear association is..."),
    answer("negative", correct = TRUE),
    answer("positive"),
    answer("relatively strong", correct = TRUE),
    answer("relatively weak"),
    post_message = "The slope of the (currently imaginary) line of best fit is negative, so the linear association is negative, and the dots cluster fairly tightly around the line of best fit, so the linear association is relatively strong.<br><br> See DSS §3.5.1.",
    random_answer_order = FALSE,
    allow_retry = TRUE
  )
)
```

```{r dis2, echo=FALSE}
dis <- read.csv("data/UK_districts.csv")
dis1 <- na.omit(dis)
```

<br> <span style=color:dodgerblue> **Q`r qnum <- qnum+1; qnum`:** </span> **Compute the correlation coefficient between *high_education* and *leave*.**

```{r ch3_topic3_r3, exercise=TRUE, exercise.setup = "dis2"}

```

<div id="ch3_topic3_r3-hint">
Hint: Use `cor()` with the two variables as arguments (in either order). See DSS §3.5.2.
</div>

```{r ch3_topic3_r3-solution}
cor(dis1$high_education, dis1$leave)
```

```{r ch3_topic3_r3-check}
grade_result(
  pass_if(~ isTRUE(all.equal(.result, cor(dis1$high_education, dis1$leave))),
          "<br><br> See DSS §3.5.2.")
)
```

```{r ch3_topic3_r3_inter}
quiz(caption=NULL,
  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, " - SELECT ALL THAT APPLY:</span> Does the correlation between *high_education* and *leave* make sense given their scatter plot? Yes, it makes sense that..."),
    answer("cor(high_education, leave) < 0, because in the scatter plot the slope of the line of best fit is negative, indicating a negative linear association", correct = TRUE),
    answer("cor(high_education, leave) < 0, because in the scatter plot the dots cluster fairly tightly around the line of best fit, indicating a strong linear association", message="The sign of the correlation coefficient reflects the direction of the linear association, not the strength."),
    answer("|cor(high_education, leave)| is closer to 1 than to 0, because in the scatter plot the slope of the line of best fit is negative, indicating a negative linear association", message = "The absolute value of the correlation coefficient reflects the strength of the linear association, not the direction."),
    answer("|cor(high_education, leave)| is closer to 1 than to 0, because in the scatter plot the dots cluster fairly tightly around the line of best fit, indicating a strong linear association", correct = TRUE),
    post_message = "The sign of the correlation reflects the direction of the association, which the slope of the line of best fit shows; its absolute value reflects the strength, which is how tightly the dots cluster around that line.<br><br> See DSS §3.5.2.",
    random_answer_order = FALSE,
    allow_retry = TRUE
  ),
  
  question(paste0("<span style=color:dodgerblue> Q", qnum <- qnum+1, ":</span> From this analysis, we conclude that, at the district level, a higher proportion of highly educated residents..."),
    answer("was typically associated with a lower vote share for Brexit", correct = TRUE),
    answer("was typically associated with a higher vote share for Brexit", message="Check the direction: the correlation is negative."),
    answer("caused lower Brexit support", message = "Correlation does not necessarily imply causation."),
    answer("caused higher Brexit support", message = "Correlation does not necessarily imply causation."),
    post_message = "Districts with more highly educated residents tended to record lower leave vote shares. We cannot conclude that education *caused* lower Brexit support: a third variable, such as the local economy, could be driving both. Here, we are analyzing data from the entire population, not from a sample, so we do not need to assume that the sample is representative. (We assume that dropping two districts with missing data does not affect this.)<br><br> See DSS §3.5.2.",
    random_answer_order = TRUE,
    allow_retry = TRUE
  )
)
```

## <span style=color:dodgerblue> GENERATE PROGRESS REPORT
You have reached the end of this tutorial. Congratulations! To generate a progress report, enter your name below and click the *Generate Report* button. Take a screenshot if your instructor asks for proof of your work.

```{r context="server"}
shiny::observeEvent(
  input$get_score, 
  {
    objs2 <- learnr:::get_tutorial_state()
    
    # Keep only the graded questions and exercises
    scored <- Filter(
      function(x) isTRUE(x$type %in% c("question", "exercise")),
      objs2
    )
    
    # Number of questions/exercises answered
    num_answered <- length(scored)
    
    # Number of questions/exercises answered correctly
    num_correct <- sum(vapply(scored, function(x) isTRUE(x$correct), logical(1)))
    
    # Formats a proportion as a percentage, guarding against dividing by zero
    pct <- function(numerator, denominator) {
      if (isTRUE(denominator > 0)) {
        paste0(round(numerator / denominator * 100, 0), "%")
      } else {
        "not available yet"
      }
    }
    
    # Wraps text in bold dodgerblue
    bold <- function(x) {
      paste0("<b>", x, "</b>")
    }
    
    output$chapter = shiny::renderUI(shiny::HTML(
      paste0("PROGRESS REPORT OF DSS: ",
             bold(toupper(as.character(chapter_id))))))
    
    output$student = shiny::renderUI(shiny::HTML(
      paste("Name:", bold(htmltools::htmlEscape(input$name)))))
    
    output$date = shiny::renderUI(shiny::HTML(
      paste0("Date and Time: ", bold(format(Sys.time(), "%Y-%m-%d")), " ",
             format(Sys.time(), "%H:%M:%S"))))
  
    output$user = shiny::renderUI(shiny::HTML(
      paste("User Name:", bold(htmltools::htmlEscape(Sys.info()[c("user")])))))
    
    output$num_answered = shiny::renderUI(shiny::HTML(
      paste0("Exercises Answered: ", bold(num_answered),
             " (out of ", num_tutorial, ")")))
    
    output$num_correct = shiny::renderUI(shiny::HTML(
      paste("Exercises Answered Correctly:", bold(num_correct))))
    
    output$completion_rate = shiny::renderUI(shiny::HTML(
      paste("Proportion of Exercises Answered:",
            bold(pct(num_answered, num_tutorial)))))
    
    output$prop_correct_answered = shiny::renderUI(shiny::HTML(
      paste("Proportion Correct Out of Answered:",
            bold(pct(num_correct, num_answered)))))
    
    output$prop_correct_all = shiny::renderUI(shiny::HTML(
      paste("Proportion Correct Out of All:",
            bold(pct(num_correct, num_tutorial)))))
    
    invisible()
  }
)
```

```{r score, echo=FALSE}
shiny::textInput("name", "Please Enter Your Name")
shiny::actionButton("get_score", "Generate Report")
shiny::br()
shiny::br()
shiny::br()
shiny::uiOutput("chapter")
shiny::br()
shiny::uiOutput("student")
shiny::uiOutput("date")
shiny::uiOutput("user")
shiny::br()
shiny::uiOutput("num_answered")
shiny::uiOutput("num_correct")
shiny::br()
shiny::uiOutput("completion_rate")
shiny::uiOutput("prop_correct_answered")
shiny::br()
shiny::uiOutput("prop_correct_all")
shiny::br()
shiny::br()
```
<p class="borderexample">
<b>This tutorial is still a work in progress — I would love your feedback!</b><br><br>
Was it helpful? Did you find anything confusing? Do you have any other comments or suggestions?<br>
Please let me know by completing <a href="https://suffolk.co1.qualtrics.com/jfe/form/SV_b4m4HmOumNWHY6W">this short survey</a>. (It can be submitted anonymously.)<br><br>
<i>Thank you, and I hope you enjoyed the tutorial!</i> – Elena
</p>

